<nav class="navbar navbar-default">
  <div class="container-fluid">
    <div class="navbar-header" style="float: left">
        <a class="navbar-brand" href="0_Index.ipynb" target="_self"> <h2> &uarr; Back to front page</h2></a>
    </div>
  </div>
</nav>

# Exploring an audio file

When audio signals are recorded to a digital medium, this is done by sampling the amplitude of an analog signal (aka. voltage output from a microphone) at regular intervals known as the *sampling period* $T_s$, resulting in a sequence of measured values. An illustration of this process is shown below.

<!--<img src="Figures/Task1_Ts.png" style="width: 600px; margin-left: 100px" />-->

![Sampling period illustrated](Figures/Task1_Ts.png)

Closely related to the sampling period $T_s$ is the *sampling rate* or *sampling frequency* $f_s$, defined as the *number of samples per second*. The sampling frequency $f_s$ is calculated simpy by inverting the samping period $T_s$.

$$f_s =\frac{1}{T_s}$$

In order to reconstruct an analog voltage signal which can be sent to speakers or headphones, both knowledge of the sampling frequency used during recording *and* the measured sample values are required. Happily, all this information is contained within any audio file format.

The codecell below uses a function available through the `scipy` module to read the audio file `example_music.wav` containing some relaxing classical music. The sequence of audio amplitude measurements which contain the actual audio signal is assigned to the variable `sampleData` in the form of a very long `array`, while the sampling frequency is an integer value which is assigned to the variable `fs`. 

After loading the audio file, the data is passed along to an embedded Auduio playback widget accessible through the IPython module. 

```python
import scipy.io.wavfile as wavfile # Import module for handling of .wav audio files
from IPython.display import Audio   # For loading embedded audio player

fs, sampleData = wavfile.read("sample_audio.wav") # "fs" is sampling frequency, "sampleData" is the sequence of measurements

# Use the following lines to listen to the audio signal
Audio(sampleData, rate=fs)
```

## a)

Run the codecell and listen to the audio clip. We can now begin our analysis of the audio signal. 

Using the data gathered in the codecell above, write a script which calculates the sampling period $T_s$ in *microseconds*, and the total duration of the audio file. <br> *hint: the function `len()` should be useful here*. 

<!-- grade_id: cell-3be1130d2ae8705f, points: 5, grade: true, solution: true -->
```python
T_s = "???" # Sample period in microseconds
duration = "???" # Audio file length in seconds

# WRITE YOUR CODE IN THIS CELL:
```

<!-- grade_id: cell-95ffc67c8bcfcb87, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(T_s, 3)).encode('utf-8')+b'e585b').hexdigest() == '86114d68d6f3bf7c408f2570aa7837f6b2918733', 'Wrong answer for T_s :('; print('Correct answer for T_s :)')
assert sha1(str(round(duration, 3)).encode('utf-8')+b'35160').hexdigest() == '6359777fd70e173df99cec0e2997b1f33c1d0065', 'Wrong answer for duration :('; print('Correct answer for duration :)')
```

## b)

Use the Audio object to play back the audio clip once with double the sample rate $f_s$, and once with half the sample rate. What effects can you hear? Explain the cause of these effects.

<!-- grade_id: cell-7938e705a223d241, points: 5, grade: true, solution: true -->
```python
# WRITE YOUR CODE IN THIS CELL:
```

<!-- grade_id: cell-6965f45985179140, points: 5, grade: true, solution: true -->

<div class="alert alert-info">
<h4> Answer theory questions here!</h4>
</div>


A crucial parameter to be aware of when working with audio is the bit rate of the digital audio signal measured in bps (bits per second).

## c) 

Given each sample having 16 bit resolution, what is the bit rate for the audio recording and what is the total size of the file (assuming no compression)?


<!-- grade_id: cell-508f57c803ebff8d, points: 5, grade: true, solution: true -->
```python
bit_rate = "???" # bits per second
file_size = "???" # Total file size in bits

# WRITE YOUR CODE IN THIS CELL:
```

<!-- grade_id: cell-c7a68e8fca8ca302, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(round(bit_rate)).encode('utf-8')+b'7031e').hexdigest() == 'e977d8c1a62c632cc44688e4d9291ad1ffc6293c', 'Wrong answer for bit_rate :('; print('Correct answer for bit_rate :)')
assert sha1(str(round(file_size)).encode('utf-8')+b'de0f5').hexdigest() == '283536d927f25736dd1699d4c53e580290d6a099', 'Wrong answer for file_size :('; print('Correct answer for file_size :)')
```

It is often desirable to provide a graphical plot of the audio signal. Simply passing the `sampleData` to the `plot()` function should acheive this, but the x-axis will show sample index `n` instead of the time `t` in seconds, which may be rather unhelpful. Based on the total number of samples in the audio clip and the sampling rate, we can adjust the plot to show signal amplitude as a function of time $t$ in seconds by creating a new array (e.g. `t`) of equal length to our audio signal which spans the duration of our audio signal.

For example, given an audio file with $N$ samples and sampling rate $f_s$, what we want is an array `t` with `N` elements linearly spaced in the time span $0 \leq t < \frac{N}{f_s}$.

## d)

Write the code to generate this plot in the codecell below. `numpy.array`, `numpy.arange` and `numpy.linspace` can all be used here. If done correctly, the resulting figure should look something like [this](Figures/audio_plot.png).


<!-- grade_id: cell-779cad4651c44182, points: 5, grade: true, solution: true -->
```python
# WRITE YOUR CODE IN THIS CELL:
```

<br>
<nav class="navbar navbar-default">
  <div class="container-fluid">
    <div class="navbar-header" style="float: right">
        <a class="navbar-brand" href="2_processing_audio.ipynb" target="_self">Next page: <i>Processing audio signals</i> &gt;</a>
    </div>
  </div>
</nav>